Fragmentation and aggregation of cyanobacterial colonies

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Taken altogether, the experimental evidence favors an erosion-dominated process. However, a few minor questions remain regarding the models. Why does the equalfragmentation model predict no biomass transfer between size classes? To what extent, quantitatively, does the erosion model outperform the equal fragments model at capturing the biomass size distributions? Finally, why does the idealized erosion fail to capture the size distribution at late stages in Supplemental Figure S9 - would this discrepancy be resolved if the authors considered individual colony variances in cell adhesion (for instance, as hypothesized by the authors in lines 133-137)? I do not believe these questions curb the other results of the paper.

Our analysis in Figure 2 considers two size classes: small colonies (l < 5) and large colonies (l ≥ 5). The equal-fragment model predicts that the fracture of a large colony gives rise to two daughter fragments with half the biovolume. For an average colony of l = 25 in diameter, this corresponds to two daughter fragments with a diameter of l = 18, which is still in the large colony class. Sequential fragmentation events would be required to set a biomass transfer to the small size range (l < 5). However, the nearly exponential behavior of the fragmentation frequency function (Eq. 5) implies that subsequent fragmentation events are greatly slowed down. Therefore, the equal-fragments model predicts that the biomass transfer from large to small colonies during the first five hours of the experiment is negligible. This is in a sharp contrast with the erosion model, which transfers biomass to the small size class at every fragmentation event. The difference between the two fragmentation models is quantified in Figure 2D, with a negligible change in biomass size distribution for the equal-fragment model (horizontal dash-dotted line) and a strong increase of small colonies for the erosion model (curved dashed line). Hence, it is clear from Figure 2D that the erosion model outperforms the equal-fragment model by capturing the observed shift from large to small colonies. We have now described this more clearly in lines 231-233.

Nevertheless, the performance of the idealized erosion model is limited at late stages (Fig. S9D). We agree with the reviewer that this limitation could potentially be overcome with the introduction of variance in cell adhesion among colonies (as we hypothesized in lines 140142). However, this is not a trivial thing to do, as it would require additional free parameters and reduce the simplicity of the model. Therefore, we chose to restrain our model to the common assumptions of idealized fragmentation models widely used in literature (e.g. references 53-55).

Reviewer #2 (Public review):

Especially the introduction seems to imply that shear force is a very important parameter controlling colony formation. However, if one looks at the results this effect is overall rather modest, especially considering the shear forces that these bacterial colonies may experience in lakes. The main conclusion seems that not shear but bacterial adhesion is the most important factor in determining colony size. The writing could have done more justice to the fact that the importance of adhesion had been described elsewhere. This being said, the same method can be used to investigate systems where shear forces are biologically more relevant.

In this work we aimed to investigate the effects of shear forces over a wide range of values, extending beyond the regime of natural lakes into the strong mixing created by technological applications such as the bubble plumes that are applied in several lakes to suppress cyanobacterial blooms. The adhesion force between cells via, e.g., extracellular polysaccharides (EPS) play an essential role by controlling the resistance to shear-driven erosion, which has been quantified in our model by the fitting parameters Si and qi.

We agree with the reviewer that we have missed some literature on Microcystis colony formation via cell aggregation (i.e., cell adhesion), for which we apologize. In our new revision, we have now included several new references [30-34,36] and we now describe the findings of these earlier studies. Specifically, in the Introduction we now pay more attention to the role of cell adhesion by writing (lines 53-60):

“In contrast, cell aggregation (sometimes also called cell adhesion) can promote a rapid increase in colony size beyond the limit set by division rates, and may explain sudden rises in colony size in late bloom periods [26, 30, 31]. Aggregation rates depend on the stickiness of the colonies, which in turn is controlled by the EPS composition, pH, and ionic composition of water [27–29]. In particular, divalent cations such as Ca2+ can bridge negatively charged functional groups in EPS and therefore increase stickiness [32–34]. It has been shown that high levels of Ca2+ enhance cell aggregation in Microcystis cultures [35]. Moreover, cell aggregation can provide a fast defense against grazing [36]. Fluid flow plays an important role in cell aggregation by regulating the collision frequency between cells or colonies [6]. In addition, fluid flow ….”

Furthermore, in the Conclusions we added (lines 374-376):

“A previous study on colony aggregation at high Ca2+ levels observed similar morphological differences in colony formation [35]. There, an initial fast cell aggregation produced a sparse colony structure, followed by a more compact structure of the colonies associated with cell division”

Finally, we would like to clarify a difference in terminology between the reviewer’s comment and our work. The term cell adhesion is commonly used in microbiology to refer to adhesion of cells with a solid substrate. In our work, the adhesion mediated by EPS occurs between free-floating cells and colonies. To avoid any confusion, we chose to refer to this process as cell aggregation, in line with other literature on suspended particles.

Reviewer #2 (Recommendations for the authors):

The authors have expanded on the image analysis process but now report substantially different correction factors (λ2 =2.79 compared to 73.13 in the previous submission; λ3 =0.52 compared to 13.71 in the previous submission). Could the authors comment on how the analysis changed? These correction factors for N<5 appear particularly relevant for the aggregation experiments presented in Figure 3. For measurements involving only small colonies, as in Figure 3, are these correction factors still valid? In addition, does the timing of image acquisition, i.e. when the colonies are imaged, influence the correction factors applied in this study?

The description of the calibration process was improved in our earlier revision of the manuscript to improve clarity and remove unclear definitions. In the first version, the supplementary equation (S1) for the input variable Np[i] was defined as the number of features per frame. This variable is dependent on the frame dimension (2048x2048 px for large colonies, l>5, and 400x400 px for small colonies). We believe that a more suitable input is the concentration distribution, which is normalized by frame area, and therefore invariant to frame dimensions and less prone to misinterpretations. For this reason, we adjusted this definition of Np[i] in the revised version of the manuscript, so that it expresses the number of features per frame area (instead of per frame). These changes required the calibration constants, λ2 and λ3, to be updated in the manuscript by a factor of (400 px/2048 px)2. This explains why these two calibration constants changed by a factor 0.038. This rescaling of the input variable Np[i] and the calibration constants did not affect the final results of our calculations (Figures 2 and 3).

The authors use a moderate dissipation rate to stir the colonies, after which they allow them to sediment. How long were the particles allowed to sediment before measurements were taken? Intuitively, one might expect a greater number of colonies to be detected following sedimentation, yet the authors report only about one third of the colonies in the sedimented state. What accounts for this reduction? Furthermore, if higher shear rates are applied, do the results differ, for instance if particles are lifted further by the shear flow? Some more clarity would help other researchers to perform similar work.

The sedimentation of particles following an initial stir was applied only for creating a reference size distribution, displayed in the supplementary Figures S8-C and D. As one intuitively would expect, a higher concentration of colonies was detected after sedimentation (Fig. S8-C and D) than during the shear flow (Fig. S8-A and B). During all other experiments in our work, the applied dissipation rate was sufficient to ensure a uniform distribution of colonies in suspension throughout the parameter range, as described in lines 461-473.

In the caption of Figure S8 we have reported the number of colonies counted in small subsamples. These numbers are just small subsets of the total number of colonies contained in the entire volume of the cone-and-plate setup. A sub-sample with larger volume was measured during the shear flow in comparison to the sub-sample measured for the sedimented sample, leading to a larger number of counted colonies in panels A and B (N = 10776, combined) compared to panels C and D (N = 3066 and 1455, respectively).

However, when normalized for the volume of the sub-samples, the calculated concentration of colonies is higher for panels C and D (as shown in the graphs). We understand that the earlier caption description of Figure S8 was misleading, for which we apologize. In the revised version, we have adjusted the caption to better describe the quantity:

“Number of colonies counted during sampling …”

Line 797 contains an unfinished edit ("Figure ADD") that should be corrected.

The unfinished edit has been corrected in the newly revised manuscript. Thanks!

Changes from the earlier version to the revised version

Abstract

Earlier version: Fluid flow has a major effect on the aggregation and fragmentation of bacterial colonies. Yet, a generic framework to understand and predict how hydrodynamics affects colony size remains elusive. This study investigates how fluid flow affects the formation and maintenance of large colonial structures in cyanobacteria, using an experimental technique that precisely controls hydrodynamic conditions. We performed experiments on laboratory cultures and lake samples of the cyanobacterium Microcystis, while their colony size distribution was measured simultaneously by direct microscopic imaging. We demonstrate that EPSembedded cells formed by cell division exhibit significant mechanical resistance to shear forces. However, at elevated hydrodynamic stress levels (exceeding those typically generated by surface wind mixing) these colonies experience fragmentation through an erosion process. We also show that single cells can aggregate into small colonies due to fluid flow. However, the structural integrity of these flow-induced colonies is weaker than that of colonies formed by cell division. We provide a mathematical analysis to support the experiments and demonstrate that a population model with two categories of colonies describes the measured size distributions. Our results shed light on the specific conditions wherein flow-induced fragmentation and aggregation of cyanobacteria are decisive and indicate that colony formation under natural conditions is mainly driven by cell division, although flow-induced aggregation could play a role in dense bloom events. These findings can be used to improve prediction models and mitigation strategies for toxic cyanobacterial blooms and also offer potential applications in other areas such as algal biotechnology or medical settings where the dynamics of biological aggregates play a significant role.

Revised version: Fluid flow has a major effect on the aggregation and fragmentation of bacterial colonies. Yet, a generic framework to understand and predict how hydrodynamics affects colony size remains elusive. This study investigates how fluid flow affects the formation and maintenance of large colonial structures in cyanobacteria, using an experimental technique that precisely controls hydrodynamic conditions. We performed experiments on laboratory cultures and lake samples of the cyanobacterium Microcystis, while their colony size distribution was measured simultaneously by direct microscopic imaging. We demonstrate that EPS-embedded cells formed by cell division exhibit significant mechanical resistance to shear forces. However, at elevated hydrodynamic stress levels (exceeding those typically generated by surface wind mixing) these colonies experience fragmentation through an erosion process. We also show that single cells can aggregate into small colonies due to fluid flow. However, the structural integrity of these flow-induced colonies is weaker than that of colonies formed by cell division. We provide a mathematical analysis to support the experiments and demonstrate that a population model with two categories of colonies describes the measured size distributions. Our results shed light on the specific conditions wherein flow-induced fragmentation and aggregation of cyanobacteria are decisive and indicate that colony formation under natural conditions is mainly driven by cell division, although flow-induced aggregation could play a role in dense bloom events. These findings can be used to improve prediction models and mitigation strategies for toxic cyanobacterial blooms and also offer potential applications in other areas such as algal biotechnology or medical settings where the dynamics of biological aggregates play a significant role.

Introduction

Many environmental and pathogenic bacteria exhibit colonial forms of life, either as surface biofilms [1, 2] or non-attached aggregates [3–5]. In most situations, these colonies are subjected to hydrodynamic stresses, which may lead to flow-induced fragmentation and aggregation [6]. These physical processes, combined with biological growth, determine the fate of aggregates and the microbial lifestyle. Nonetheless, despite their abundance and fundamental importance, a generic framework to study the effects of fluid flow on bacterial colonies, as well as their fragmentation and aggregation across scales, is still in its infancy [7–10]. Here, we focus on colonies of bloom-forming cyanobacteria and ask the question: how flow-induced aggregation and fragmentation regulate dynamic changes in colony size?

Toxic cyanobacterial blooms are a major nuisance for freshwater ecosystems worldwide [11]. These blooms often lead to increased turbidity, cause harm to other aquatic organisms, and deteriorate water quality. Recent evidence indicates that the severity and frequency of cyanobacterial blooms are increasing due to eutrophication and climate change [12–15]. Consequently, water management authorities show growing interest in techniques to prevent, monitor, and control these toxic blooms [16]. A particular effort is being taken to combat blooms of Microcystis, a non-motile colony-forming freshwater cyanobacterial genus that dominates many aquatic ecosystems worldwide [17, 18]. Fluid dynamics is at the heart of Microcystis bloom formation, but is also used in some of the methods to control them. A large colony size and positive buoyancy give Microcystis a higher average light dose in comparison to non-buoyant algae [19, 20]. Control methods such as deep artificial mixing have been successfully applied in several lakes to disrupt this competitive advantage by increasing the turbulence level [21]. Yet, the mechanisms of Microcystis colony formation, and in particular, the impact of fluid flow on the formation of colonies are poorly understood.

Added in the revised version: Microcystis cells are held together by a layer of extracellular polymeric substances (EPS), facilitating the formation of large colonies of up to a millimeter in diameter with diverse morphological arrangements [22–24].

Earlier version: Microcystis cells are held together by a layer of extracellular polymeric substances (EPS), facilitating the formation of large colonies of up to a millimeter in diameter with diverse morphological arrangements [22–24]. Two main mechanisms for colony formation have been proposed [17]: (i) cell division, i.e., clonal expansion of cells which remain attached by their secreted EPS layer after division, and (ii) aggregation, i.e., initially independent cells/colonies collide and adhere to each other. Sequencing of Microcystis colonies has revealed a much higher genetic diversity between independent colonies than within colonies, thus supporting the hypothesis that colony formation under natural conditions is primarily driven by cell division [25]. However, cell aggregation is a possible route for increase in colony size on a short time scale, and may offer improved protection against high grazing pressure [26]. Aggregation becomes significant when the collision frequency is high enough for aggregation rates to overcome cell division rate [6], and the stickiness of colonies is sufficiently large (controlled by the EPS composition, pH , and ionic strength of water or medium [27–29]). Yet, it remains an open question to what extent aggregation of Microcystis cells and colonies is a relevant process in natural systems. In addition to aggregation, fluid flow also influences colony formation through fragmentation, potentially serving as a mechanism for colony reproduction [30] to overcome the growth limitations of large colonies [31]. Beyond modulating colony size through aggregation and fragmentation, fluid flow has also been proposed as a driving factor in changes in colony morphology and growth rate [32–34].

Revised version: Two main mechanisms for colony formation have been proposed [17]: (i) cell division, i.e., clonal expansion of cells which remain attached by their secreted EPS layer after division, and (ii) aggregation, i.e., initially independent cells/colonies collide and adhere to each other. Sequencing of Microcystis colonies has revealed a much higher genetic diversity between independent colonies than within colonies, thus supporting the hypothesis that colony formation under natural conditions is primarily driven by cell division [25]. In contrast, cell aggregation (sometimes also called cell adhesion) can promote a rapid increase in colony size beyond the limit set by division rates, and may explain sudden rises in colony size in late bloom periods [26–28]. Aggregation rates depend on the stickiness of the colonies, which in turn is controlled by the EPS composition, pH, and ionic composition of water [29–31]. In particular, divalent cations such as Ca2+ can bridge negatively charged functional groups in EPS and therefore increase stickiness [32–34]. It has been shown that high levels of Ca2+ enhance cell aggregation in Microcystis cultures [35]. Moreover, cell aggregation can provide a fast defense against grazing [36]. Fluid flow plays an important role in cell aggregation by regulating the collision frequency between cells or colonies [6]. In addition, fluid flow also influences colony formation through fragmentation, potentially serving as a mechanism for colony reproduction [37] to overcome the growth limitations of large colonies [38]. Beyond modulating colony size through aggregation and fragmentation, fluid flow has also been proposed as a driving factor in changes in colony morphology and growth rate [39–41].

Earlier version: Flow-induced alterations in colony size and morphology impact the vertical migration and grazing resistance of colonies, thereby influencing the success of the Microcystis population. This bears significance for forecasting models which have often assumed a fixed colony size [35–37]. Information on colony size is also important for bloom control via artificial mixing systems, which rely on estimates of colony flotation velocity, determined by both colony shape and size [21, 38]. Moreover, it remains uncertain whether mixing systems, such as bubble plumes, can enhance their efficiency by fragmenting colonies or whether artificial mixing may inadvertently promote aggregation. Colony fragmentation has been observed in grid-stirred [39] and propeller mixed tanks [32] at energy dissipation rates comparable to bubble plumes [40]. However, these studies reported averaged dissipation rates, while local values near the fast moving surfaces can be orders of magnitude higher [41]. Furthermore, the size distribution of the colony fragments has not been characterized.

Revised version: Flow-induced alterations in colony size and morphology impact the vertical migration and grazing resistance of colonies, thereby influencing the success of the Microcystis population. This bears significance for forecasting models which have often assumed a fixed colony size [42–44]. Information on colony size is also important for bloom control via artificial mixing systems, which rely on estimates of colony flotation velocity, determined by both colony shape and size [21, 45]. Moreover, it remains uncertain whether mixing systems, such as bubble plumes, can enhance their efficiency by fragmenting colonies or whether artificial mixing may inadvertently promote aggregation. Colony fragmentation has been observed in grid-stirred [46] and propeller mixed tanks [39] at energy dissipation rates comparable to bubble plumes [47]. However, these studies reported averaged dissipation rates, while local values near the fast moving surfaces can be orders of magnitude higher [48]. Furthermore, the size distribution of the colony fragments has not been characterized.

In the present study, we investigated to what extent Microcystis colony fragmentation and/or aggregation occurs at levels of turbulence observed for naturally or artificially mixed lakes. To this end, colonies of a laboratory culture of Microcystis were subjected to various intensities of shear flow (Fig. 1). We identified the mechanism of colony fragmentation using direct measurements of dynamic changes in colony size distribution under controlled laboratory conditions and analyze the characteristics of colonies formed by aggregation and cell division. Furthermore, the fragmentation behavior of the laboratory culture was compared with field samples of Microcystis spp. We also developed a numerical model based on a two-category discrete population to simulate changes in the size of Microcystis colonies via aggregation and fragmentation. Finally, we provide a phase map based on flow characteristics and cyanobacterial abundance to highlight different regimes in which fluid flow plays a significant role in cyanobacterial colony formation.

Methodology used to observe effects of fluid flow on Microcystis colonies.

(A) Experimental setup consisting of a (cone-and-plate) controlled flow setup combined with inverted microscopy. The conical upper surface was rotated by a rheometer head, while the stationary glass slide below the sample allowed optical access for the microscope. (B) Examples of microscopy images. Colony size distributions were calculated after image processing of the captured frames. (C) Changes in size distributions (and other complementary measurements) over time were used to identify aggregation and fragmentation of cyanobacterial colonies. (D) The majority of the measurements were conducted using a laboratory culture of Microcystis strain V163. Colonies collected from Lake Gaasperplas (Netherlands), dominated by the morphospecies M. aeruginosa were also used.

Fragmentation of Large Microcystis Colonies occurs through Erosion

A culture of Microcystis strain V163 was filtered to collect large colonies, and colonies on the filter were subsequently re-dispersed in fresh BG-11 medium up to the desired total biovolume fraction ϕ (Materials and Methods). The flow was generated by a cone-and-plate shear operating in an inertial regime, and it was characterized by the average energy dissipation rate ε˙ , hereafter called dissipation rate (Materials and Methods). The size distribution of the colonies was measured over time by microscopic imaging of the suspension through the transparent lower surface (Fig. 1). The colony size is described here by its relative diameter l = d/d1, where d is the Feret’s diameter (maximum distance between two boundary points of the colony) and d1 is the single-cell diameter. The number of cells in a colony, Nc, was calculated from its relative diameter via the relation Nc = lf, where f = 2.09 is the average fractal dimension measured for colonies of Microcystis strain V163 (Materials and Methods). The size distribution was computed as the normalized biovolume distribution n(l), i.e., the number of cells in colonies in size bin {l; l + Δ l}, divided by the total number of cells in the suspension and the bin width.

The results show that the initial size distribution had a bimodal shape, with a subpopulation of small colonies peaking around l ~ 1 — 2 (single cells, dimers, and trimers) and a subpopulation of large colonies, peaking around l ~ 20 — 40 (Fig. 2A). The cut-off size between the two subpopulations is defined here as the minimum of the bimodal distribution. Some small colonies (l ~ 1 — 2), up to 20% of the total biovolume, were always present in the initial distribution even after filtration, due to pre-stirring required to keep the suspension homogeneous. However, medium-sized colonies (l = 4 — 8) accounted for a negligible fraction of the distribution.

Added in the revised version: Kinetics of fragmentation of Microcystis strain V163 colonies under cone-and-plate shear flow.

Added in the revised version: The laboratory culture was filtered to select mainly large colonies, and the total biovolume fraction was adjusted to ϕ = 10-4. Suspensions were subjected to an intense dissipation rate ( ε˙=5.8m2/s3 ) in panels A-E. (A) The initial size distribution of colonies, expressed as biovolume fraction of the relative colony diameter (normalized by single cell diameter). The size distribution had a bimodal shape, with large colonies (l > 5, in yellow) and small colonies (l ≤ 5, in green, composed mostly of single cells, dimers, and some trimers, as depicted by the inset). (B) Median diameter of small and large colonies as a function of time. Symbols indicate the experimental data, while the lines indicate the predictions from the population model given by Eq. (1) with an erosion (dashed) or equal fragments (dot-dash) hypothesis for large division-formed colonies. Small aggregates colonies follow the equal fragments hypothesis. Bars and shaded region indicate limits of 25th and 75th percentiles. (C) Most large colonies have been fragmented after 1 hour of shear flow, but the bimodal shape of the colony size distribution remained. (D) Biovolume fraction of small colonies (i.e., biovolume of small colonies over the total biovolume) as a function of time. A shift is observed from large to small colonies, captured well by the erosion model, but not by the equal fragments model. (E) The rate of change in biovolume distribution at t = 0h. Negative values indicate loss of colonies by fragmentation, while positive values indicate newly created fragments. The distribution suggests an erosion mechanism, as depicted by the cartoon inside the plot. (F) Fragmentation frequency as a function of the relative diameter of colonies for three values of dissipation rate. The inset shows details for moderate dissipation rates. Error bars indicate the standard deviation and the lines indicate the predictions by Eq. (5). Best fit parameters: α1 = 0.023, S1 = 0.034 , q1 = 4.5; erosion: S2 = 31 , q2 = 4. 1 ; equal fragments: S2 = 33 , q2 = 6.3.

When subjected to an intense dissipation rate ( ε˙=5.8m2/s3 ), the large colonies of Microcystis strain V163 displayed considerable fragmentation, with their median diameter (50 % percentile of the biovolume distribution of the large colonies) decaying to about half of the initial value after a few hours (Fig. 2B). In contrast, the median diameter of small colonies showed little change in size and remained below l ≤ 2 during the entire experiment. The size distribution after 1 hour indicated that most colonies larger than l ≳ 30 were fragmented, leading to a rise in the subpopulation of small colonies (Fig. 2C). Notably, the bimodal shape of the size distribution was preserved, with very few medium-sized colonies. This suggests that aggregation of single cells into larger colonies is negligible under this intense shear flow condition, and the main effect of the shear flow on the size distribution is the decrease in size of the large colonies concomitant with an increase in the fraction of small colonies. After a few hours of shear, about 80 % of the total biovolume was composed of small colonies (Fig. 2D). This time scale was much shorter than the cell division time, such that changes in the colony size distribution by division could be neglected within the experimental time [20].

Removed in the revised version: Kinetics of fragmentation of Microcystis strain V163 colonies under cone-and-plate shear flow. The laboratory culture was filtered to select mainly large colonies, and the total biovolume fraction was adjusted to ϕ = 10−4. Suspensions were subjected to an intense dissipation rate ( ε˙=5.8 m2/s3 ) in panels A-E.

Removed in the revised version: (A) The initial size distribution of colonies, expressed as biovolume fraction of the relative colony diameter (normalized by single cell diameter). The size distribution had a bimodal shape, with large colonies (in yellow) and small colonies (in green, composed mostly of single cells, dimers, and some trimers, as depicted by the inset). (B) Median diameter of small and large colonies as a function of time. Symbols indicate the experimental data, while the lines indicate the predictions from the population model given by Eq. (1) with an erosion (dashed) or equal fragments (dot-dash) hypothesis for large division-formed colonies. Small aggregates colonies follow the equal fragments hypothesis. Bars and shaded region indicate limits of 25th and 75th percentiles. (C) Most large colonies have been fragmented after 1 hour of shear flow, but the bimodal shape of the colony size distribution remained. (D) Biovolume fraction of small colonies (i.e., biovolume of small colonies over the total biovolume) as a function of time. A shift is observed from large to small colonies, captured well by the erosion model, but not by the equal fragments model. (E) The rate of change in biovolume distribution at t = 0h. Negative values indicate loss of colonies by fragmentation, while positive values indicate newly created fragments. The distribution suggests an erosion mechanism, as depicted by the cartoon inside the plot. (F) Fragmentation frequency as a function of the relative diameter of colonies for three values of dissipation rate. The inset shows details for moderate dissipation rates. Error bars indicate the standard deviation and the lines indicate the predictions by Eq. (5). Best fit parameters: α1 = 0.023, S1 = 0.034, q1 = 4.5; erosion: S2 = 31, q2 = 4.1; equal fragments: S2 = 33, q2 = 6.3.

Earlier version: Direct measurement of the fragment distribution function g(l,l′) (i.e., the frequency of fragments with size l arising from the fragmentation of a colony of size l′) is an experimental challenge [42–44]. Individual fragmentation events could not be visualized as the residence time of colonies inside the field of view is too short. Instead, a suitable model can be identified from the rate of change in the biovolume distribution (Fig. 2E). A negative rate was observed for large colonies in the range l ≳ 25, indicating a loss in that size range by fragmentation. In contrast, the newly created fragments appeared as positive peaks, with a distinct peak close to the single-cell size. This result is compatible with an erosion mechanism, in which small fragments with one or a few cells are detached from the outer layer of the colony. For an ideal binary erosion, a fragmentation event of a colony with i cells results in one cell being eroded from the surface of the colony with a complementary fragment of i − 1 cells remaining. Small deviations from an ideal binary erosion are observed in Figure 2E, with most eroded fragments lying in the range 1≤ l ≤ 2. Nevertheless, ideal erosion provides a good approximation for the results. Accurate identification of complementary fragments from the rate of change in the biovolume distribution was more difficult, because the broad initial distribution of our suspension led to a broad distribution of complementary fragments. In this case, the rate of change in the biovolume distribution shifted from a net loss of eroded colonies at larger sizes (l ≳ 25) to a net gain of complementary fragments at intermediate sizes (5 ≲ l ≲ 20).

Revised version: Direct measurement of the fragment distribution function g(l,l′) (i.e., the frequency of fragments with size l arising from the fragmentation of a colony of size l′) is an experimental challenge [49–51]. Individual fragmentation events could not be visualized as the residence time of colonies inside the field of view is too short. Instead, a suitable model can be identified from the rate of change in the biovolume distribution (Fig. 2E). A negative rate was observed for large colonies in the range l ≳; 25, indicating a loss in that size range by fragmentation. In contrast, the newly created fragments appeared as positive peaks, with a distinct peak close to the single-cell size. This result is compatible with an erosion mechanism, in which small fragments with one or a few cells are detached from the outer layer of the colony. For an ideal binary erosion, a fragmentation event of a colony with i cells results in one cell being eroded from the surface of the colony with a complementary fragment of i — 1 cells remaining. Small deviations from an ideal binary erosion are observed in Figure 2E, with most eroded fragments lying in the range 1 ≤ l ≤ 2. Nevertheless, ideal erosion provides a good approximation for the results. Accurate identification of complementary fragments from the rate of change in the biovolume distribution was more difficult, because the broad initial distribution of our suspension led to a broad distribution of complementary fragments. In this case, the rate of change in the biovolume distribution shifted from a net loss of eroded colonies at larger sizes (l ≳ 25) to a net gain of complementary fragments at intermediate sizes (5 ≲ l ≲ 20).

Assuming an ideal binary erosion mechanism, the fragmentation frequency κ (l) was computed from the temporal dependence of the biovolume distribution (Materials and Methods). Here, κ(l) is defined as the average number of fragmentation events (loss of a single cell) that one colony of size l suffers per unit of time. The fragmentation frequency of Microcystis strain V163 was calculated as a function of colony size for three values of dissipation rate (Fig. 2F). The results show that the fragmentation frequency increases with colony size. Furthermore, an increase in dissipation rate sharply intensifies the erosion of colonies, with a two order of magnitude increase in κ(l) when ε˙ varies from 1.1 to 9.6 m2/s3. The intensified fragmentation frequency also reduces the time required for the biovolume distribution to reach a steady-state (SI Appendix, Fig. S3). It is also noticeable that some of the large colonies (5 to 10 % of the total biovolume) do not fragment to single cells but remain as large colonies of 10-20 cells even after 5 hours of very high dissipation rate ( ε˙=9.6m2/s3 ), indicating variability in the mechanical strength of colonies, with a subset presenting an exceptional resistance in comparison to the average behavior (SI Appendix, Fig. S3E and F)

Aggregation-Formed Colonies are Less Resistant than Division-Formed Colonies

In addition to inducing fragmentation, a turbulent flow may also increase the collision frequency of colonies. If the collisions are adhesive and the suspension is concentrated enough, then new colonies will be formed by aggregation. We subjected a single-cell suspension to the cone-and-plate shear flow to investigate whether aggregation can be a dominant colony-forming mechanism compared to cell division. Cultures of the Microcystis strain V163 were filtered to obtain a single-cell suspension, and the total biovolume fraction was adjusted with fresh BG-11 medium (Materials and Methods). A typical initial bio-volume fraction distribution is presented in Figure 3A. In addition to single cells, some small colonies up to l = 3 were present in the filtered suspension due to imperfect filtering and/or aggregation during the pre-stirring. The size distribution was unimodal (single peak, Fig. 3A), in contrast with the bimodal distribution (double peak, Fig. 2A) in the suspension of division-formed colonies. After about 1 hour at a moderate dissipation rate ( ε˙=0.019m2/s3 ) and a total biovolume fraction of ϕ = 10-4, the single-cell suspension showed aggregation of cells into small colonies, with a steady-state median diameter of l = 1.8 (Figs. 3B and 3C). Note that cell division can be neglected as the measurement time was much shorter than the division time. An increase in the collision rates due to a denser single-cell suspension led to an increase in the median diameter, although the increase was minor. At a high total biovolume fraction (ϕ = 5.10-4), a steady-state median diameter of l = 2.4 was reached, with a negligible fraction of colonies above l ≳ 5 (SI Appendix, Fig. S4C). Therefore, the aggregation of single cells is likely limited by weak bonds between cells. At the steady-state regime of median diameter, the newly-formed aggregates with a large diameter would be quickly fragmented back into smaller parts.

Earlier version: Kinetics of aggregation for a single-cell suspension of Microcystis strain V163 at a moderate dissipation rate of ε˙=0.019m2/s3 and a total biovolume fraction of ϕ = 10−4.

Revised version: Kinetics of aggregation for a single-cell suspension of Microcystis strain V163 at a moderate dissipation rate of ε˙ = 0.019 m2/s3 and a total biovolume fraction of ϕ = 10-4.

(A) Initial size distribution of the suspension as a function of the relative diameter, composed mostly of single cells. (B) Median diameter of colonies formed by aggregation of single cells as a function of time. Symbols indicate the experimental data, while the lines indicate the predictions from the population model given by Eq. (1) with an equal fragment hypothesis for small aggregates colonies. Bars and shaded region indicate limits of 25th and 75th percentiles. (C) After 1 hour of shear flow, the size distribution has shifted towards slightly larger diameters. (D) Time behavior of a suspension of large division-formed colonies under the same moderate dissipation rate and total biovolume fraction. The bimodal size distribution is separated into large (yellow) and small (green) colonies. Dashed lines and shaded regions indicate predictions from the population model given by Eq. (1). Best fit parameters: α1 = 0.023, S1 = 0.034 , q1 =4.5, S2 =31 , q2 = 4.1.

A suspension of division-formed colonies under the same moderate dissipation rate ( ε˙=0.019m2/s3 ) did not suffer noticeable fragmentation (Fig. 3D). The culture was grown in quiescent medium, without shaking, thus the colonies are assumed to have formed solely by cell division. The erosion of division-formed colonies was negligible, as the fraction of small colonies remained below 20% (SI Appendix, Fig. S2B). Furthermore, the nearly constant median diameter of the large colonies (observed also for other total biovolume fraction values, SI Appendix, Fig. S2) shows that the division-formed colonies did not aggregate at this moderate dissipation rate.

The results above indicate that bonds formed during flow-induced aggregation (whether between two single cells or two division-formed colonies) are weak and cannot withstand intense dissipation rates. In contrast, the bonds between cells within a single division-formed colony are significantly more resistant. This difference in mechanical strength of the cellular binding in aggregation-formed versus division-formed colonies can be attributed to the structure of the EPS layer that surrounds the colonies. Aggregated colonies would have intercellular gaps (smaller that the cell size) that are not filled with EPS. In contrast, cell division allows enough time for the gaps to be filled with secreted EPS, effectively increasing the bond strength between cells (Figs. 3B and 3D). However, we could not confirm this hypothesis visually, as the EPS layer in Microcystis strain V163 is too thin (< 1 μm) for a clear identification using optical microscopy.

Dynamical Changes in Colony Size Modeled by a Two-Category Distribution

Earlier version: Our experimental results for colonies of Microcystis strain V163 (Fig. 3) suggest the presence of two categories of colonies: single cells and small aggregates (denoted as C1) and large colonies formed by cell division (C2). The total biovolume distribution n(l) is given by the sum of both distributions n1(l)+n2(l) , normalized such that ∫0∞n(l)dl=1 . Within times shorter than the cell division time, the biovolume distribution of each colony category is described by a population model [27]: ∂ni∂t(l,t)=Ai(l,t)+Fi(l,t)+Ti(l,t), (1)

Revised version: Our experimental results for colonies of Microcystis strain V163 (Fig. 3) suggest the presence of two categories of colonies: single cells and small aggregates (denoted as C1) and large colonies formed by cell division (C2). The total biovolume distribution n(l) is given by the sum of both distributions n1 (l) + n2(l), normalized such that ∫0∞n(l)dl=1 . Within times shorter than the cell division time, the biovolume distribution of each colony category is described by a population model [29]: ∂ni∂t(l,t)=Ai(l,t)+Fi(l,t)+Ti(l,t), (1)

where Ai(l, t) and Fi(l, t) describe the aggregation and fragmentation rates as a function of colony size for each of the two categories (i ∈ [1, 2]), and Ti(l,t) is a transfer rate of biovolume between the two categories. The turbulent shear-induced collisions between colonies shift the distribution towards larger aggregates, and the aggregation rate is given by Ai(l,t)=6ϕαiπ[12∫0lβ(l′,lo)(ll′lo)2f−1ni(l′,t)ni(lo,t)dl′ −ni(l,t)∫0∞β(l,l′)1l′fni(l′,t)dl′] , (2)

Earlier version: where lo=(lf−l′f)1/f . The collision frequency is described by the kernel β(l,l′) and for a turbulent shear flow, we have β(l,l′)=0.163(l+l′)3ε˙/ν , where ν is kinematic viscosity of the fluid [45]. The ratio of aggregative collisions over total collisions is given by the stickiness parameter αi. For the category C1 (single cells and small aggregates), the stickiness is positive, α1 > 0. In contrast, we neglect the aggregation of C2 colonies (collisions between C2 − C2 and C1 − C2), since the bonds formed by aggregation would be weaker and quickly rupture again, such that α2 = 0. This assumption is corroborated by the absence of aggregation between large colonies in our experiments (Fig. 3D).

Revised version: where lo = (lf — l′f)1 /f. The collision frequency is described by the kernel β(l, l′) and for a turbulent shear flow, we have β(l,l′)=0.163(l+l′)3ε˙/ν , where ν is kinematic viscosity of the fluid [52]. The ratio of aggregative collisions over total collisions is given by the stickiness parameter αi. For the category C 1 (single cells and small aggregates), the stickiness is positive, α1 > 0. In contrast, we neglect the aggregation of C2 colonies (collisions between C2-C2 and C1 -C2), since the bonds formed by aggregation would be weaker and quickly rupture again, such that α2 = 0. This assumption is corroborated by the absence of aggregation between large colonies in our experiments (Fig. 3D).

The colony formation will be limited by fragmentation if the colony exceeds a critical size li* , given by li*=(ε˙Si)−qi/f, (3)

Earlier version: where Si and qi are constitutive parameters that depend on the aggregate structure and bond strength, with a lower bound at li∗≥1 [46, 47]. We assume that aggregation-formed colonies have a small critical size, as the overlap between the EPS layer of each cell is narrow, leading to weak bonds. Meanwhile, division-formed colonies produce sufficient EPS to fill the space between cells, thus increasing the resistance to hydrodynamic stress. The fragmentation rate is given by a first-order kinetic equation [46]: Fi(l,t)=−κi(l)ni(l,t)+∫l∞gi(l,l′)κi(l′)ni(l′,t)dl′ (4)

Revised version: where Si and qi are constitutive parameters that depend on the aggregate structure and bond strength, with a lower bound at li*≥1 [53, 54]. We assume that aggregation-formed colonies have a small critical size, as the overlap between the EPS layer of each cell is narrow, leading to weak bonds. Meanwhile, division-formed colonies produce sufficient EPS to fill the space between cells, thus increasing the resistance to hydrodynamic stress. The fragmentation rate is given by a first-order kinetic equation [53]: Fi(l,t)=−κi(l)ni(l,t)+∫l∞gi(l,l′)κi(l′)ni(l′,t)dl′, (4)

Earlier version: where κi(l) is the fragmentation frequency, which is a function of the probability distribution function of the hydrodynamic stress. Assuming a Gaussian distribution, the fragmentation frequency is [46]: κi(l)=2ε˙15πνexp⁡(−Mi2)erf(Mi), (5)

Revised version: where κi (l) is the fragmentation frequency, which is a function of the probability distribution function of the hydrodynamic stress. Assuming a Gaussian distribution, the fragmentation frequency is [53]: κi(l)=2ε˙15πνexp(−Mi2)erf(Mi), (5)

Earlier version: where Mi=15/2(l/li∗)−f/2qi . Due to the scarcity of experimental data in the literature, simplified models are usually assumed for the fragment distribution function gi(l,l′), such as erosion, equal fragments, or a uniform distribution [48]. Assuming a binary fragmentation (i.e., two fragments per fragmentation event), the fragment size distribution can be written in the form gi(l,l′)=δ(l−la)+δ(l−lb) , where la and lb are the sizes of fragments. Based on the results from division-formed colonies under intense dissipation rate (Fig. 2E), an ideal erosion mechanism is adopted for category C2 colonies, such that g2(l,l′)=δ(l−1)+δ(l−l′f−1f) . In contrast, the experimental results for the small colonies (Fig. 3C) show a wider size distribution of eroded fragments which can not be captured by an ideal erosion mechanism, so an equal fragments mechanism is adopted for category C1 colonies, i.e., g2(l,l′)=2δ(l−l′/2f) . The size distributions of C1 and C2 colonies are coupled through a transfer rate of biovolume Ti(l,t). We propose that the cells eroded from C2 colonies are transferred to category C1, such that: T1(l,t)=−T2(l,t)=δ(l−1)∫l∞κ2(l′)n2(l′,t)dl′ (6)

Revised version: where Mi=15/2(l/li*)−f/2qi . Due to the scarcity of experimental data in the literature, simplified models are usually assumed for the fragment distribution function gi(l,l′), such as erosion, equal fragments, or a uniform distribution [55]. Assuming a binary fragmentation (i.e., two fragments per fragmentation event), the fragment size distribution can be written in the form gi(l,l′) = δ(l — la) + δ(l — lb), where la and lb are the sizes of fragments. Based on the results from division-formed colonies under intense dissipation rate (Fig. 2E), an ideal erosion mechanism is adopted for category C2 colonies, such that g2(l,l′)=δ(l−1)+δ(l−l′f−1f) . In contrast, the experimental results for the small colonies (Fig. 3C) show a wider size distribution of eroded fragments which can not be captured by an ideal erosion mechanism, so an equal fragments mechanism is adopted for category C1 colonies, i.e., g2(l,l′)=2δ(l−l′/2f) . The size distributions of C1 and C2 colonies are coupled through a transfer rate of biovolume Ti (l, t). We propose that the cells eroded from C2 colonies are transferred to category C1 , such that: T1(l,t)=−T2(l,t)=δ(l−1)∫l∞κ2(l′)n2(l′,t)dl′, (6)

while the complementary fragments of C2 colonies remain in that category.

Earlier version: We solved the balance equation (1) using a zero-order discrete population balance, or sectional method, in which the size distribution is described by uniformly sized bins where the size distribution is evaluated at a representative size centered on the bin [48, 49]. The simulations were initialized with the size distributions at t = 0 obtained in our experiments. The two pairs of parameters for critical size ( {q1,S1} and {q2,S2} ) and the stickiness parameter (α1) were used to fit the model to the experimental data (Materials and Methods).

Revised version: We solved the balance equation (1) using a zero-order discrete population balance, or sectional method, in which the size distribution is described by uniformly sized bins where the size distribution is evaluated at a representative size centered on the bin [55, 56]. The simulations were initialized with the size distributions at t = 0 obtained in our experiments. The two pairs of parameters for critical size ({q1 ,S1} and {q2,S2}) and the stickiness parameter (α1) were used to fit the model to the experimental data (Materials and Methods).

Earlier version: For an initial suspension composed mainly of large colonies at an intense dissipation rate ( ε˙=5.8m2/s3 ), the model prediction of a decrease in the median size of C2 colonies up to a steady-state diameter is in good agreement with the experimental findings (Fig. 2B). Moreover, the erosion mechanism accurately predicts the biovolume transfer from large to small colonies (Fig. 2D). In comparison, the hypothesis of equal fragments formation for C2 colonies does not capture such transfer, as all fragments remain in the large colony range (Fig. 2D and Fig. S9 of SI Appendix). The use of the erosion mechanism for the fragment distribution function is further corroborated by previous studies simulating bacterial aggregates [50]. For C1 colonies at this intense dissipation rate, the model predicts that their median diameter remains near l∗ = 1 (i.e., single cells). However, the experimental results show a slightly wider size distribution of both the small and large colonies than predicted by an ideal erosion mechanism (Fig. 2C). Nevertheless, the global behavior of the size distribution is well described by the two-category model. Furthermore, a very good agreement is also observed between the measured fragmentation frequency and the model predictions given by Eq. (5) (Fig. 2F).

Revised version: For an initial suspension composed mainly of large colonies at an intense dissipation rate ( ε˙=5.8m2/s3 ), the model prediction of a decrease in the median size of C2 colonies up to a steady-state diameter is in good agreement with the experimental findings (Fig. 2B). Moreover, the erosion mechanism accurately predicts the biovolume transfer from large to small colonies (Fig. 2D). In comparison, the hypothesis of equal fragments formation for C2 colonies predicts that all fragments remain within the large colony size category (l > 5) during the first 5 hours of experiments, hence the equal fragments model does not capture the observed transfer from large to small colonies (Fig. 2D and Fig. S9 of SI Appendix).The use of the erosion mechanism for the fragment distribution function is further corroborated by previous studies simulating bacterial aggregates [57]. For C1 colonies at this intense dissipation rate, the model predicts that their median diameter remains near l* = 1 (i.e., single cells). However, the experimental results show a slightly wider size distribution of both the small and large colonies than predicted by an ideal erosion mechanism (Fig. 2C). Nevertheless, the global behavior of the size distribution is well described by the two-category model. Furthermore, a very good agreement is also observed between the measured fragmentation frequency and the model predictions given by Eq. (5) (Fig. 2F).

Earlier version: To complement the mathematical model and to further investigate the mechanical origins of the resistance of division-formed colonies to intense shear, we performed rheology tests on the cyanobacterial colonies (Materials and Methods). Simple and oscillatory shear measurements revealed elasto-viscoplastic mechanics [51, 52] of the concentrated aggregate with a yield stress of τy=4.3±0.3Pa . The low volume fraction of cells in aggregate (∼ 12 %) indicates that the mechanical properties are mainly dictated by the EPS layer. Equating the average hydrodynamic shear stress to the yield stress in Eq. (11) and subsequently substituting the critical dissipation rate in Eq. (3) results in a critical colony size of l2∗=2.8 . The eroded fragments observed in the fragmentation experiments (Fig. 2E) are mostly smaller than this critical size. This suggests that large colonies will suffer erosion when the local hydrodynamic stress exceeds the yield stress of the EPS layer. Although the fracture behavior of colonies in suspension is expected to be different from a concentrated aggregate under shear rheology, the latter provides a reasonable and convenient estimate of the material properties of the colonies. Related to the critical size of the colony and their mechanical properties is the exponent q2 (see equation 5). The best fit of the results in Fig. 2 provided q2 = 4.1. This value lies in the upper limit of previous results for colloidal aggregates [42, 46]. Note that for colloids, q is related to the stress concentration around cracks in the fractal structure of the aggregate [47], and low values of q are associated with brittle deformation, in which the stress concentration grows quickly with aggregate size. In contrast, cyanobacterial colonies have a soft plastic behavior that is less sensitive to large cracks, thus justifying the large value of q.

Revised version: To complement the mathematical model and to further investigate the mechanical origins of the resistance of division-formed colonies to intense shear, we performed rheology tests on the cyanobacterial colonies (Materials and Methods). Simple and oscillatory shear measurements revealed elasto-viscoplastic mechanics [58, 59] of the concentrated aggregate with a yield stress of τy = 4.3 ± 0.3 Pa. The low volume fraction of cells in aggregate (~ 12 %) indicates that the mechanical properties are mainly dictated by the EPS layer. Equating the average hydrodynamic shear stress to the yield stress in Eq. (11) and subsequently substituting the critical dissipation rate in Eq. (3) results in a critical colony size of l2* = 2.8. The eroded fragments observed in the fragmentation experiments (Fig. 2E) are mostly smaller than this critical size. This suggests that large colonies will suffer erosion when the local hydrodynamic stress exceeds the yield stress of the EPS layer. Although the fracture behavior of colonies in suspension is expected to be different from a concentrated aggregate under shear rheology, the latter provides a reasonable and convenient estimate of the material properties of the colonies. Related to the critical size of the colony and their mechanical properties is the exponent q2 (see equation 5). The best fit of the results in Fig. 2 provided q2 = 4.1. This value lies in the upper limit of previous results for colloidal aggregates [49, 53]. Note that for colloids, q is related to the stress concentration around cracks in the fractal structure of the aggregate [54], and low values of q are associated with brittle deformation, in which the stress concentration grows quickly with aggregate size. In contrast, cyanobacterial colonies have a soft plastic behavior that is less sensitive to large cracks, thus justifying the large value of q.

We further tested the model with the aggregation experiments. A suspension of C1 colonies with an initial distribution close to the single-cell size (Fig. 3A) was simulated with an equal fragments hypothesis. The model agrees well with the experimental results for the median colony size, accurately predicting a steady-state value when there is a balance between collisions and fragmentation (Fig. 3B). The best fit of the results provided a small value for the stickiness parameter (α1 = 0.023). The aggregation rate is linearly dependent on the stickiness and the total biovolume fraction (Eq. 2), thus an increase in these parameters (α1 and ϕ) allows the suspension to reach the steady-state size faster. However, the steady-state size has low sensitivity to these parameters (α1 and ϕ), as the fragmentation frequency has a higher-order dependence with colony size (Fig. 2F). The equal fragments hypothesis leads to a single-peaked size distribution, whose steady-state median increases with the total biovolume fraction. If an ideal erosion mechanism was assumed for the C1 colonies, the peak would be fixed around the single-cell size (SI Appendix, Fig. S6). The experimental results for the C1 colonies are in better agreement with the equal fragments hypothesis (Fig. 3C).

Earlier version: Note that the ideal erosion and equal fragments mechanisms are idealized scenarios, and the real distribution is expected to be more dispersed. Furthermore, a fractal description of small colonies is limited, as a continuum function between biovolume and diameter of colonies may not hold. Moreover, the turbulence model used for fragmentation (Eq. 5) poses two limitations: (i) the theory was developed for homogeneous Gaussian turbulence [46], which differs considerably from the boundary-layer turbulence generated in a cone-and-plate shear [53], and (ii) the colonies were larger than the Kolmogorov scale for the intense dissipation rate used in our fragmentation measurements (η = 20 μm for a dissipation rate of ε˙=5.8m2/s3 , SI Appendix), meaning that the colony size is comparable to the smallest eddies in the flow. Despite these limitations, the proposed two-category model provides a good approximation of the fragmentation process observed in our experiments. Simulations that take into account spatially inhomogeneous and non-Gaussian turbulence are demanding and lack an explicit analytic expression for κ(l) that can be applied to a population model [54].

Revised version: Note that the ideal erosion and equal fragments mechanisms are idealized scenarios, and the real distribution is expected to be more dispersed. Furthermore, a fractal description of small colonies is limited, as a continuum function between biovolume and diameter of colonies may not hold. Moreover, the turbulence model used for fragmentation (Eq. 5) poses two limitations: (i) the theory was developed for homogeneous Gaussian turbulence [53], which differs considerably from the boundary-layer turbulence generated in a cone-and-plate shear [60], and (ii) the colonies were larger than the Kolmogorov scale for the intense dissipation rate used in our fragmentation measurements (η = 20 μm for a dissipation rate of ε˙=5.8m2/s3 , SI Appendix), meaning that the colony size is comparable to the smallest eddies in the flow. Despite these limitations, the proposed two-category model provides a good approximation of the fragmentation process observed in our experiments. Simulations that take into account spatially inhomogeneous and non-Gaussian turbulence are demanding and lack an explicit analytic expression for κ (l) that can be applied to a population model [61].

Field Samples of Microcystis are more Resistant to Fragmentation Due to Thicker EPS Layer

Samples of Microcystis spp. were collected from Lake Gaasperplas - the Netherlands (Materials and Methods). A major difference in fragmentation behavior is observed between the laboratory cultures of Microcystis strain V163 and the field samples of Microcystis spp. While the laboratory cultures suffered intense fragmentation at a dissipation rate of ε˙=5.8m2/s3 , the colonies collected from the field showed negligible fragmentation at this dissipation rate (Fig. 4A). The size distribution of field colonies also had a bimodal shape both initially and after 1 hour of flow (Figures 4C and E). The small colonies displayed an increase in size due to aggregation up to 2.5 times the initial median size (Fig. 4D), although the fraction of small colonies remained below 11% even after a few hours of shear (Fig. 4F). The large colonies displayed only a small increase in size due to aggregation and quickly reached a steady-state size.

Kinetics of the fragmentation of colonies in field samples of Microcystis spp. at an intense dissipation rate ( ε˙=5.8m2/s3 ) and total biovolume fraction of ϕ = 10-4.

(A) Comparison of the fragmentation frequency as a function of colony size for the laboratory culture (Microcystis strain V163) and the field samples (Microcystis spp.). Error bars indicate the standard deviation. (B) Brightfield microscopy images of colonies in a Nigrosin-dyed medium (dark region) show evidence of a thick EPS layer (bright region) surrounding a field colony. (C) Initial size distribution of colonies in field samples as a function of the relative diameter. The size distribution had a bimodal shape, with small colonies (green) and large colonies (yellow). (D) The median diameter of the colonies in each subpopulation as a function of time. Bars indicate 25th and 75th percentiles. (E) After 1 hour of shear flow, the small colonies have aggregated slightly, while the large colonies keept their size distribution. (F) The biovolume fraction of small colonies remained nearly constant during the experiment.

Earlier version: The higher resistance of field colonies to shear can be attributed to the thick EPS layer that surrounds the colony, which could be seen under brightfield microscopy using a negative staining (Fig. 4B, Materials and Methods). In contrast, this layer was barely visible in colonies of laboratory cultures of Microcystis strain V163. Nevertheless, the limited amount of aggregation observed for large field colonies indicates that aggregationformed colonies have a smaller resistance to shear than division-formed colonies, similar to the results observed for laboratory colonies (Fig. 3). Differences in the morphology of colonies from field samples and laboratory cultures are frequently observed in the literature, with the latter often presenting only single cells or small colonies [55]. Colony formation may be influenced by various factors [17], where especially high calcium and low nitrogen concentrations are identified as factors promoting high EPS production [56–58]. This is consistent with the calcium and nitrogen concentrations measured in Lake Gaasperplas (71 mg/L Ca2+ and 0.04 mg/L N) versus BG-11 medium (9.8 mg/L Ca2+ and 246 mg/L N) (Materials and Methods).

Revised version: The higher resistance of field colonies to shear can be attributed to the thick EPS layer that surrounds the colony, which could be seen under brightfield microscopy using a negative staining (Fig. 4B, Materials and Methods). In contrast, this layer was barely visible in colonies of laboratory cultures of Microcystis strain V163. Nevertheless, the limited amount of aggregation observed for large field colonies indicates that aggregation-formed colonies have a smaller resistance to shear than division-formed colonies, similar to the results observed for laboratory colonies (Fig. 3). Differences in the morphology of colonies from field samples and laboratory cultures are frequently observed in the literature, with the latter often presenting only single cells or small colonies [62]. Colony formation may be influenced by various factors [17], where especially high calcium and low nitrogen concentrations are identified as factors promoting high EPS production [35, 63, 64]. This is consistent with the calcium and nitrogen concentrations measured in Lake Gaasperplas (71 mg/L Ca2+ and 0.04 mg/L N) versus BG-11 medium (9.8 mg/L Ca2+ and 246 mg/L N) (Materials and Methods).

Aggregation, Fragmentation, or Cell Division? A Guideline

Earlier version: Under which conditions are aggregation, fragmentation or cell division the dominant processes in colony formation? In a quiescent fluid, cyanobacterial colonies are expected to grow solely due to cell division with a specific growth rate μ=Nc−1dNc/dt , which is a function of various environmental factors such as light intensity and nutrient availability [17, 20]. Let l¯=∫0∞ln(l,t)dl be the average colony size of a suspension. Assuming that the cells remain attached after division, the rate of increase in the average colony size will be 1l¯dl¯dt=μf (7)

Revised version: Under which conditions are aggregation, fragmentation or cell division the dominant processes in colony formation? In a quiescent fluid, cyanobacterial colonies are expected to grow solely due to cell division with a specific growth rate μ=Nc−1dNc/dt , which is a function of various environmental factors such as light intensity and nutrient availability [17, 20]. Let l¯=∫0∞l n(l,t)dl be the average colony size of a suspension. Assuming that the cells remain attached after division, the rate of increase in the average colony size will be 1l¯dl¯dt=μf. (7)

Earlier version: The increase in colony size will progress until a maximum size lm is reached. The upper limit in colony size is still a subject of intense discussion, with species variability, reduced division rate at large sizes, and dissolution of the EPS layer being important factors [17, 31]. If the suspension is concentrated enough, collisions between colonies will also increase the average colony size due to aggregation. Let us consider initially a monodisperse suspension with a single size l¯ . At short times, the rate of increase in the average colony size due to aggregation will be (SI Appendix) 1l¯dl¯dt=2.49(21/f−1)ϕαl¯4−2fε˙ν (8)

Revised version: The increase in colony size will progress until a maximum size lm is reached. The upper limit in colony size is still a subject of intense discussion, with species variability, reduced division rate at large sizes, and dissolution of the EPS layer being important factors [17, 38]. If the suspension is concentrated enough, collisions between colonies will also increase the average colony size due to aggregation. Let us consider initially a monodisperse suspension with a single size l¯ . At short times, the rate of increase in the average colony size due to aggregation will be (SI Appendix) 1l¯dl¯dt=2.49(21/f−1)ϕαl¯4−2fε˙ν. (8)

Assuming f = 2 for simplicity (which is similar to the measured value for Microcystis strain V163), we can use Eqs. (7) and (8) to define a dimensionless number Ag, quantifying the ratio between the rates of increase in the average colony size due to aggregation and cell division, given by Ag=2.06αϕμε˙ν (9)

Earlier version: The dimensionless ratio Ag reveals the dominant mechanism for colony formation (cell division is dominant for Ag ≪ 1, while aggregation is dominant for Ag ≫ 1) and depends on the total biovolume fraction ϕ and dissipation rate ε˙ , as shown by the phase diagram in Figure 5. For a lake with a total biovolume fraction corresponding to an early bloom ( ϕ=3⋅10−7 , according to WHO alert level 1 [59]), a low dissipation rate typical of wind-induced turbulence ( ε˙∼10−7m2/s3 at a depth of ∼ 1 m) [60], a specific growth rate of μ ∼ 0.008 h−1 [20], a kinematic viscosity of ν=10−6m2s−1 and a stickiness of α ∼ 0.02, we estimate a ratio Ag ∼ 10−3. Under these conditions, the colony size increases due to cell division while the role of aggregation is negligible (region I in Fig. 5). In contrast, during a dense bloom event, the scum layer can reach a very high biovolume fraction (ϕ ∼ 10−2 [61, 62]) in the top layer (depth of ∼ 10−2 m), where concomitantly the dissipation rate is much higher ( ε˙∼10−4m2/s3 under weak wind turbulence [63]), which leads to Ag ∼ 105 (keeping all remaining parameters constant). Although this regime (region II in Fig. 5) was not explored in our cone-and-plate experiments, recent mesocosm experiments found clear evidence for the aggregation of Microcystis colonies in surface scums, with rapid increase in colony size at rates exceeding cell division rates [64]. Specifically, for a bulk biovolume fraction of ϕ ∼ 10−4 and a dissipation rate of ε˙∼10−4m2/s3 , that study measured a rate of increase in the average colony size of l¯−1dl¯/dt∼0.2h−1 , corresponding to a ratio of Ag ∼ 101. These results lie well inside region II of Fig. 5. These mesocosm experiments confirm our model prediction that aggregation is the dominant mechanism for colony formation in region II.

Revised version: The dimensionless ratio Ag reveals the dominant mechanism for colony formation (cell division is dominant for Ag ≪ 1, while aggregation is dominant for Ag ≫ 1) and depends on the total biovolume fraction ϕ and dissipation rate ε˙ , as shown by the phase diagram in Figure 5. For a lake with a total biovolume fraction corresponding to an early bloom (ϕ = 3 · 10-7, according to WHO alert level 1 [65]), a low dissipation rate typical of wind-induced turbulence ( ε˙∼10−7m2/s3 at a depth of ~ 1m) [66], a specific growth rate of μ~ 0.008h-1 [20], a kinematic viscosity of ν = 10-6m2s-1 and a stickiness of α~ 0.02, we estimate a ratio Ag~ 10-3. Under these conditions, the colony size increases due to cell division while the role of aggregation is negligible (region I in Fig. 5). In contrast, during a dense bloom event, the scum layer can reach a very high biovolume fraction (ϕ~ 10-2 [67, 68]) in the top layer (depth of ~ 10 -2 m), where concomitantly the dissipation rate is much higher ( ε˙∼10−4m2/s3 under weak wind turbulence [69]), which leads to Ag~ 105 (keeping all remaining parameters constant). Although this regime (region II in Fig. 5) was not explored in our cone-and-plate experiments, recent mesocosm experiments found clear evidence for the aggregation of Microcystis colonies in surface scums, with rapid increase in colony size at rates exceeding cell division rates [70]. Specifically, for a bulk biovolume fraction of ϕ~ 10-4 and a dissipation rate of ε˙∼10−4m2/s3 , that study measured a rate of increase in the average colony size of l¯−1dl¯/dt∼0.2h−1 , corresponding to a ratio of Ag~ 101. These results lie well inside region II of Fig. 5. These mesocosm experiments confirm our model prediction that aggregation is the dominant mechanism for colony formation in region II.

Phase diagram indicating the dominant colony formation mechanism as a function of the dissipation rate ε˙ and the cyanobacterial abundance (expressed by the total biovolume fraction ϕ).

Earlier version: (I) Colonies grow only by cell division at low dissipation rates and total biovolume fractions. (II) As the biovolume fraction increases, aggregation enhances colony growth. (III) For moderate dissipation rates, aggregated colonies are fragmented, and only cell division can increase colony size. (IV) Fragmentation of colonies dominates at intense dissipation rates, irrespective of whether these colonies were formed by aggregation or cell division. Bars on the right side indicate typical values of dissipation rate observed for natural wind-mixing, bubble plumes in artificially mixed lakes and laboratory-scale setups such as cone-and-plate systems and stirred tanks. The dashed arrowed line indicates the transition from an early bloom (left bullet - WHO alert level 1 [59]) to a dense surface scum (right bullet - typical scum biovolume fraction [61]) under typical wind mixing.

Revised version: (I) Colonies grow only by cell division at low dissipation rates and total biovolume fractions. (II) As the biovolume fraction increases, aggregation enhances colony growth. (III) For moderate dissipation rates, aggregated colonies are fragmented, and only cell division can increase colony size. (IV) Fragmentation of colonies dominates at intense dissipation rates, irrespective of whether these colonies were formed by aggregation or cell division. Bars on the right side indicate typical values of dissipation rate observed for natural wind-mixing, bubble plumes in artificially mixed lakes and laboratory-scale setups such as cone-and-plate systems and stirred tanks. The dashed arrowed line indicates the transition from an early bloom (left bullet - WHO alert level 1 [65]) to a dense surface scum (right bullet - typical scum biovolume fraction [67]) under typical wind mixing.

Earlier version: Fragmentation limits the maximum colony size to a critical value l∗. Our experimental results (Fig. 5) indicate that aggregated colonies have a lower critical size than division-formed colonies for the same turbulence intensity. Conversely, a higher critical dissipation rate is needed for division-formed colonies ( ε˙div ) than for aggregated colonies ( ε˙agg ) to achieve the same critical size l∗ = lm. This creates an interesting regime at high biovolume fractions and dissipation rates in the range ε˙agg<ε˙<ε˙div , in which only division-formed colonies can survive the turbulent flow, even though aggregation is faster than division (Ag≫1) (region III in Fig. 5). Typical values of critical dissipation rate for our experiments are ε˙agg=2.1⋅10−3m2/s3 and ε˙div=3.7m2/s3 , assuming a maximum colony size of lm = 50. Under natural turbulent conditions (e.g., surface wind-mixing), the dissipation rate ranges from ε˙∼10−8m2/s3 [60] to peaks of ε˙∼10−3m2/s3 at strong winds (in the absence of wave breaking), which is still insufficient to induce colony fragmentation [63]. Artificial mixing systems can achieve levels of turbulence that are many orders of magnitude higher. In the core of bubble plumes, dissipation rates of ε˙∼10−4−10−1m2/s3 are reached [40]. This value was sufficient to prevent aggregation of large colonies in our experiments but was not high enough to induce fragmentation of division-formed colonies. Moreover, colonies remain only a short time (∼ 1 min) in the bubble plume core, after which they are dispersed to weaker turbulent regions (dissipation rates of ε˙∼10−7−10−4m2/s3 just a few meters away from the bubble plume core). Nevertheless, bubble plumes can prevent the formation of a scum layer not only by fragmentation of newly formed weak aggregates, but also by providing sufficient turbulent mixing to disperse the buoyant colonies, thereby reducing the local biovolume fraction at the surface [21]. The intense hydrodynamic stress required to fragment division-formed colonies (region IV in Fig. 5) can only be achieved by mixers with fast-moving surfaces, such as the cone-and-plate setup used in our experiments ( ε˙∼10−2−101m2/s3 ) or propellers used in stirred tanks ( ε˙∼100−102m2/s3 ) [41,65]. These mixing devices can provide intense dissipation rates in small fluid volumes, but their implementation across entire lakes would be excessively energy-consuming.

Revised version: Fragmentation limits the maximum colony size to a critical value l*. Our experimental results (Fig. 5) indicate that aggregated colonies have a lower critical size than division-formed colonies for the same turbulence intensity. Conversely, a higher critical dissipation rate is needed for division-formed colonies ( ε˙div ) than for aggregated colonies ( ε˙agg ) to achieve the same critical size l* = lm. This creates an interesting regime at high biovolume fractions and dissipation rates in the range ε˙agg<ε˙<ε˙div , in which only division-formed colonies can survive the turbulent flow, even though aggregation is faster than division (Ag ≫ 1) (region III in Fig. 5). Typical values of critical dissipation rate for our experiments are ε˙agg=2.1⋅10−3m2/s3 and ε˙div=3.7m2/s3 , assuming a maximum colony size of lm = 50. Under natural turbulent conditions (e.g., surface wind-mixing), the dissipation rate ranges from ε˙∼10−8m2/s3 [66] to peaks of ε˙∼10−3m2/s3 at strong winds (in the absence of wave breaking), which is still insufficient to induce colony fragmentation [69]. Artificial mixing systems can achieve levels of turbulence that are many orders of magnitude higher. In the core of bubble plumes, dissipation rates of ε˙∼10−4−10−1m2/s3 are reached [47]. This value was sufficient to prevent aggregation of large colonies in our experiments but was not high enough to induce fragmentation of division-formed colonies. Moreover, colonies remain only a short time (~ 1 min) in the bubble plume core, after which they are dispersed to weaker turbulent regions (dissipation rates of ε˙∼10−7−10−4m2/s3 just a few meters away from the bubble plume core). Nevertheless, bubble plumes can prevent the formation of a scum layer not only by fragmentation of newly formed weak aggregates, but also by providing sufficient turbulent mixing to disperse the buoyant colonies, thereby reducing the local biovolume fraction at the surface [21]. The intense hydrodynamic stress required to fragment division-formed colonies (region IV in Fig. 5) can only be achieved by mixers with fast-moving surfaces, such as the cone-and-plate setup used in our experiments ( ε˙∼10−2−101m2/s3 ) or propellers used in stirred tanks ( ε˙∼100−102m2/s3 ) [48, 71]. These mixing devices can provide intense dissipation rates in small fluid volumes, but their implementation across entire lakes would be excessively energy-consuming.

Conclusions

Earlier version: Despite the widespread presence of bacterial colonies across species and environmental habitats, the influence of fluid flows on colony formation remains a fertile ground for continued exploration. A significant challenge is the lack of a generic framework in which the aggregation and fragmentation of colonies can be studied under various hydrodynamic stresses. We take a major step towards overcoming this challenge by presenting a combination of experiments and theory to elucidate the underlying mechanisms that determine the size distribution of cyanobacterial colonies over a wide range of flow intensities. Our experiments show that divisionformed Microcystis colonies can be fragmented under intense dissipation rates ε˙∼3−10m2/s3 , following an erosion mechanism. Previous studies [32, 39] have observed significant fragmentation of Microcystis colonies in field samples at dissipation rates as low as ε˙∼10−4m2/s3 . However, such studies primarily report volume averages of the dissipation rate, overlooking the likelihood of fragmentation occurring near fast-moving solid surfaces (grid stirrer [39] and propeller [32]), where the local hydrodynamic stresses are orders of magnitude higher. Instead, our results indicate that fragmentation of division-formed colonies can not be achieved with wind mixing or aeration systems.

Revised version: Despite the widespread presence of bacterial colonies across species and environmental habitats, the influence of fluid flows on colony formation remains a fertile ground for continued exploration. A significant challenge is the lack of a generic framework in which the aggregation and fragmentation of colonies can be studied under various hydrodynamic stresses. We take a major step towards overcoming this challenge by presenting a combination of experiments and theory to elucidate the underlying mechanisms that determine the size distribution of cyanobacterial colonies over a wide range of flow intensities. Our experiments show that division-formed Microcystis colonies can be fragmented under intense dissipation rates ε˙∼3−10m2/s3 , following an erosion mechanism. Previous studies [39, 46] have observed significant fragmentation of Microcystis colonies in field samples at dissipation rates as low as ε˙∼10−4m2/s3 . However, such studies primarily report volume averages of the dissipation rate, overlooking the likelihood of fragmentation occurring near fast-moving solid surfaces (grid stirrer [46] and propeller [39]), where the local hydrodynamic stresses are orders of magnitude higher. Instead, our results indicate that fragmentation of division-formed colonies can not be achieved with wind mixing or aeration systems.

Earlier version: We have demonstrated that flow-induced aggregated colonies of Microcystis form weaker structures than divisionformed colonies, likely due to the weak EPS bonds between cells. Furthermore, we showed that colony aggregation is unlikely in lakes under early bloom conditions, as the collision frequency of colonies is too low. Our results, based on the mechanical properties of colonies, are corroborated by previous studies based on the sequencing of individual Microcystis colonies [25, 66]. These studies have identified that the genetic diversity within Microcystis colonies is much smaller than the diversity between Microcystis colonies, thus supporting the hypothesis that colony formation is dominated by cell division rather than by aggregation of colonies. Nonetheless, our model supports recent studies which demonstrated that, during scum formation in very dense blooms or phytoplankton patches, flow-induced aggregation likely plays a role and may quickly increase colony size [61, 62]. A summary of these outcomes is presented in the phase map illustrated in Figure 5.

Revised version: We have demonstrated that flow-induced aggregated colonies of Microcystis form weaker structures than division-formed colonies, likely due to the weak EPS bonds between cells. A previous study on colony aggregation at high Ca2+ levels observed similar morphological differences in colony formation [35]. There, an initial fast cell aggregation produced a sparse colony structure, followed by a more compact structure of the colonies associated with cell division. Furthermore, we showed that colony aggregation is unlikely in lakes under early bloom conditions, as the collision frequency of colonies is too low. Our results, based on the mechanical properties of colonies, are corroborated by previous studies based on the sequencing of individual Microcystis colonies [25,72]. These studies have identified that the genetic diversity within Microcystis colonies is much smaller than the diversity between Microcystis colonies, thus supporting the hypothesis that colony formation is dominated by cell division rather than by aggregation of colonies. Nonetheless, our model supports recent studies which demonstrated that, during scum formation in very dense blooms or phytoplankton patches, flow-induced aggregation likely plays a role and may quickly increase colony size [67, 68]. A summary of these outcomes is presented in the phase map illustrated in Figure 5.

Earlier version: From a broader perspective, the technological advance of combining microscopic imaging with rheometry has relevance beyond predicting and mitigating cyanobacterial blooms. Since fluid flow affects the colonies and aggregates of many species [6–8, 67–70], the experimental methods and theoretical framework presented here, along with the identified mechanisms, could serve as a foundation to better understand the relationship between fluid mechanics and biological aggregates, while keeping a keen eye on the wide variation of biological functions and morphodynamics present in biological aggregates exposed to hydrodynamic stresses. An interesting example is a recent study on the fragmentation of marine snow consisting of diatoms and microplastics, which used similar methods as in our study [70, 71]. In contrast to the high mechanical resistance of Microcystis colonies, these marine snow aggregates displayed a low critical dissipation rate. Furthermore, their fragmentation was not mediated by an erosion mechanism, but by neck formation (i.e., deformation is higher in a small cross-sectional area in the middle of the aggregate). Other examples include non-natural settings, such as algae or cyanobacteria cultured in bioreactors, where dense cell aggregates experience intense mixing [72–74]. The methods and results presented here could serve as design guidelines for the study of these applications.

Revised version: From a broader perspective, the technological advance of combining microscopic imaging with rheometry has relevance beyond predicting and mitigating cyanobacterial blooms. Since fluid flow affects the colonies and aggregates of many species [6–8, 73–76], the experimental methods and theoretical framework presented here, along with the identified mechanisms, could serve as a foundation to better understand the relationship between fluid mechanics and biological aggregates, while keeping a keen eye on the wide variation of biological functions and morphodynamics present in biological aggregates exposed to hydrodynamic stresses. An interesting example is a recent study on the fragmentation of marine snow consisting of diatoms and microplastics, which used similar methods as in our study [76, 77]. In contrast to the high mechanical resistance of Microcystis colonies, these marine snow aggregates displayed a low critical dissipation rate. Furthermore, their fragmentation was not mediated by an erosion mechanism, but by neck formation (i.e., deformation is higher in a small cross-sectional area in the middle of the aggregate). Other examples include non-natural settings, such as algae or cyanobacteria cultured in bioreactors, where dense cell aggregates experience intense mixing [78–80]. The methods and results presented here could serve as design guidelines for the study of these applications.

Materials and Methods

A. Sample preparation

Earlier version: Microcystis strain V163 cultured under laboratory conditions, as well as samples collected at Lake Gaasperplas The Netherlands containing Microcystis spp. colonies were used for the experiments (Fig. 1). Microcystis strain V163 was originally isolated from Lake Volkerak - The Netherlands [75], grows in colonies and was obtained from the algal collection of the Institute for Biodiversity and Ecosystem Dynamics, University of Amsterdam. The Microcystis strain V163 was batch-cultured in Erlenmeyer flasks with standard BG-11 medium. The cultures were kept in an incubation chamber at a temperature of 20 °C. The light intensity was kept at 12 μmol photon m−2s−1 with a light:dark cycle of 14h:10h. The cultures were diluted with new BG-11 medium every week and kept at a maximum total biovolume fraction of ϕ=5⋅10−4 . The cultures were not shaken except after the addition of medium and before filtration. The Microcystis strain V163 has a single-cell diameter of d1=6.8±1.4μm (SD, N = 189). We categorize partially divided cells also as single cells, as these also behave as single units in a fragmentation/aggregation process. When considering only spherical cells not undergoing division, a diameter of 5.1 ± 0.8 μm (SD, N = 328) was obtained. Colonies of up to 500 μm were present in the culture under the incubation conditions as described before.

Revised version: Microcystis strain V163 cultured under laboratory conditions, as well as samples collected at Lake Gaasperplas - The Netherlands containing Microcystis spp. colonies were used for the experiments (Fig. 1). Microcystis strain V163 was originally isolated from Lake Volkerak - The Netherlands [81], grows in colonies and was obtained from the algal collection of the Institute for Biodiversity and Ecosystem Dynamics, University of Amsterdam. The Microcystis strain V163 was batch-cultured in Erlenmeyer flasks with standard BG-11 medium. The cultures were kept in an incubation chamber at a temperature of 20 °C. The light intensity was kept at 12 μmol photon m-2s-1 with a light:dark cycle of 14h:10h. The cultures were diluted with new BG-11 medium every week and kept at a maximum total biovolume fraction of ϕ = 5 · 10-4. The cultures were not shaken except after the addition of medium and before filtration. The Microcystis strain V163 has a single-cell diameter of d1 = 6.8 ± 1.4 μm (SD, N = 189). We categorize partially divided cells also as single cells, as these also behave as single units in a fragmentation/aggregation process. When considering only spherical cells not undergoing division, a diameter of 5.1 ± 0.8 μm (SD, N = 328) was obtained. Colonies of up to 500 μm were present in the culture under the incubation conditions as described before.

Samples of Microcystis spp. were collected from Lake Gaasperplas - the Netherlands (SI Appendix for water monitoring data and location of sampling point) on 14th August 2022. Water from the surface layer was filtered over a plankton net with a 100 μm mesh size. Samples were kept refrigerated at 4 °C, and the experiments were conducted over the next three days. The plankton composition was dominated by Microcystis spp. (89 ± 7% , SD, N = 231), with Microcystis aeruginosa being the main morphospecies, with an average single-cell diameter of d1 = 7.0± 1.4 μm (SD, N = 100). Colonies of Dolichospermum spp. were present in lower numbers. However, no differentiation of colony species was made during image processing, as it was not feasible to isolate them.

Negative staining was used to verify for the presence of an EPS layer surrounding the colonies of Microcystis strain V163 and of Microcystis spp. in field samples. Nigrosin was added to the surrounding medium, causing the EPS layer to appear bright under brightfield microscopy.

Earlier version: Upon appropriate filtering of the culture, we obtained a suspension with a specific size range. Filter paper with a 25 μm pore size was used to obtain a suspension primarily consisting of single cells (colonies with 2 or 3 cells also present in small numbers). For a suspension mainly comprising large colonies, we filtered the culture with 53 μm nylon mesh. The colonies collected in the mesh were then redispersed in a new BG-11 medium. To estimate the total biovolume fraction of the suspension, 1 ml of the sample underwent a heat treatment to disaggregate colonies [76]. At least 1000 cells were counted in a Neubauer hemocytometer assisted by the image processing software ImageJ. The total biovolume fraction of the suspension was calculated as ϕ = cv1, where c is the cell concentration and v1 is the average single-cell volume. Prior to each experiment, the desired biovolume fraction was adjusted using new BG-11 medium for the Microcystis strain V163 culture and with filtered (0.8μm) water from Lake Gaasperplas for field samples of Microcystis spp..

Revised version: Upon appropriate filtering of the culture, we obtained a suspension with a specific size range. Filter paper with a 25 μm pore size was used to obtain a suspension primarily consisting of single cells (colonies with 2 or 3 cells also present in small numbers). For a suspension mainly comprising large colonies, we filtered the culture with 53 μm nylon mesh. The colonies collected in the mesh were then redispersed in a new BG-11 medium. To estimate the total biovolume fraction of the suspension, 1 ml of the sample underwent a heat treatment to disaggregate colonies [82]. At least 1000 cells were counted in a Neubauer hemocytometer assisted by the image processing software ImageJ. The total biovolume fraction of the suspension was calculated as ϕ = cυ1 , where c is the cell concentration and υ1 is the average single-cell volume. Prior to each experiment, the desired biovolume fraction was adjusted using new BG-11 medium for the Microcystis strain V163 culture and with filtered (0.8μm) water from Lake Gaasperplas for field samples of Microcystis spp..

B. Cone-and-plate shear

Earlier version: A cone-and-plate shear flow was selected for our experiments (Fig. 1B and Fig. S7A). This geometry is commonly used for standard rheological measurements because it provides a homogeneous shear rate under a laminar flow regime. Although less common in inertial regimes, it has been recommended by previous studies to investigate the effects of hydrodynamic stress in tissues [53, 77, 78]. Furthermore, the mixing provided by turbulence helps maintain a homogeneous concentration throughout the suspension.

Revised version: A cone-and-plate shear flow was selected for our experiments (Fig. 1B and Fig. S7A). This geometry is commonly used for standard rheological measurements because it provides a homogeneous shear rate under a laminar flow regime. Although less common in inertial regimes, it has been recommended by previous studies to investigate the effects of hydrodynamic stress in tissues [60, 83, 84]. Furthermore, the mixing provided by turbulence helps maintain a homogeneous concentration throughout the suspension.

Earlier version: Our setup consists of an upper conical probe made from transparent acrylic, with an outer diameter of 42 mm and a cone complementary angle of θ = 5.44° (Fig. S7A). The conical probe is attached to a steel shaft, and it is driven by a rheometer head Anton Paar DSR 502. The lower surface consists of a 1 mm thick glass slide. A cylindrical acrylic ring confines the fluid, preventing spillage of the sample due to strong centrifugal forces. The inner diameter of the ring is 46 mm , which ensures that the nominal shear rate between the upper probe’s outer surface and the ring is the same as between the conical and lower surfaces. A clearance gap of 10 μm was set between the tip of the cone and the lower surface to avoid solid-solid friction.

Revised version: Our setup consists of an upper conical probe made from transparent acrylic, with an outer diameter of 42 mm and a cone complementary angle of θ = 5.44° (Fig. S7A). The conical probe is attached to a steel shaft, and it is driven by a rheometer head Anton Paar DSR 502. The lower surface consists of a 1mm thick glass slide. A cylindrical acrylic ring confines the fluid, preventing spillage of the sample due to strong centrifugal forces. The inner diameter of the ring is 46 mm, which ensures that the nominal shear rate between the upper probe’s outer surface and the ring is the same as between the conical and lower surfaces. A clearance gap of 10 μm was set between the tip of the cone and the lower surface to avoid solid-solid friction.

At low rotational speeds, the flow is laminar with a homogeneous shear rate γ˙n=ω/θ , where ω is the angular velocity of the conical probe. The flow transitions to a turbulent regime at faster rotations, as the inertial forces become dominant. Since the shear rate is not uniform at the non-linear regime, the nominal shear rate γ˙n is no longer an appropriate measure for shear intensity. Instead, the average intensity is given by the energy dissipation rate ε˙ which measures the rate of energy loss by viscous dissipation per unit fluid mass. Under steady-state conditions, this dissipation rate balances the power injected into the fluid, therefore: ε˙=TωρV, (10)

Earlier version: where ρ and V are, respectively, the density and volume of the fluid sample, and T is the torque on the conical probe, measured by the rheometer head. The flow regime can be identified from the energy dissipation rate curve against the angular velocity, with a a ∼ ω2 dependence in the laminar regime and a ∼ ω5/2 dependence in the turbulent regime (SI Appendix, Fig. S7B). All shear experiments conducted in this work were in the latter regime. The average hydrodynamic shear stress can be calculated from the dissipation rate by τ¯t=ρνε˙ (11)

Revised version: where ρ and V are, respectively, the density and volume of the fluid sample, and T is the torque on the conical probe, measured by the rheometer head. The flow regime can be identified from the energy dissipation rate curve against the angular velocity, with a ~ ω2 dependence in the laminar regime and a ~ ω5/2 dependence in the turbulent regime (SI Appendix, Fig. S7B). All shear experiments conducted in this work were in the latter regime. The average hydrodynamic shear stress can be calculated from the dissipation rate by T¯t=ρνε˙. (11)

Earlier version: Our cone-and-plate setup covers energy dissipation rates over three orders of magnitude, from 10 −2 to 10 m2/s3. At the lower end, the dissipation rate is limited by the sedimentation of colonies near the stagnation point at the conical tip. Due to the small gap, this stagnation point is not accessible to the colonies. A large colony of l ∼ 50 does not fit within the surfaces up to a radial distance from the tip of 4 mm . For the lowest dissipation rate (ε˙=0.019m2/s3) , a minimum secondary flow speed of ∼ 100μm/s is present [79]. This value is above the sedimentation speed of colonies of Microcystis strain V163, which are either buoyant or slightly sinking due to the presence of gas vesicles. Furthermore, the region near the tip accounts for a small fraction of sample volume, while the velocity fluctuations in the bulk region ensure that the colonies are homogeneously distributed. The upper limit for the operating dissipation rate is set by the spillage of fluid out of the cylindrical confinement. Regarding the optimal range of biovolume fraction, this is limited at ϕ = 0.0005, above which the optical overlap between nearby colonies influences the colony count. The kinematic viscosity of the liquid ν can be taken equal to water for the biovolume fraction explored here.

Revised version: Our cone-and-plate setup covers energy dissipation rates over three orders of magnitude, from 10-2 to 10 m2/s3. At the lower end, the dissipation rate is limited by the sedimentation of colonies near the stagnation point at the conical tip. Due to the small gap, this stagnation point is not accessible to the colonies. A large colony of l ~ 50 does not fit within the surfaces up to a radial distance from the tip of 4 mm. For the lowest dissipation rate ( ε˙=0.019m2/s3 ), a minimum secondary flow speed of ~ 100μm/s is present [85]. This value is above the sedimentation speed of colonies of Microcystis strain V163, which are either buoyant or slightly sinking due to the presence of gas vesicles. Furthermore, the region near the conical tip accounts for a small fraction of sample volume, while the velocity fluctuations in the bulk region of the cone-and-plate setup ensure that the colonies remain in suspension and are homogeneously distributed. The upper limit for the operating dissipation rate is set by the spillage of fluid out of the cylindrical confinement. Regarding the optimal range of biovolume fraction, this is limited at ϕ = 0.0005, above which the optical overlap between nearby colonies influences the colony count. The kinematic viscosity of the liquid ν can be taken equal to water for the biovolume fraction explored here.

The experimental setup provides a configuration where the dissipation rate can be controlled over a large range, while stress values can be measured with high precision. This, combined with optical access and simultaneous microscopy, presents a well-suited tool to study the response of biological samples under fluid flow. Our setup configuration was optimized to cover the dissipation rates observed for deep artificial mixing systems. However, the conical probe diameter could be increased and the rotation speed reduced simultaneously to cover a lower range of dissipation rate typical of wind mixing. Furthermore, the gap between the surfaces can be easily adjusted to fit other organisms.

C. Fragmentation and aggregation experiments

The fragmentation and aggregation experiments involved subjecting a suspension of Microcystis colonies to the cone-and-plate shear and using brightfield microscopy to image the colonies under shear. The lower surface was positioned on top of an inverted microscope Nikon TS100 with a 4x objective lens. The center of the field of view was located 15 mm from the rotation axis and with the focal plane set 0.5 mm above the lower surface. Illumination was provided by a Schott KL 1500 HAL light source equipped with a cyan filter (low pass with a cut-off wavelength of 555 nm) to minimize the photosynthetic activity of the sample. Images were captured with a Prime BSI Express sCMOS camera with a resolution of 1.63 μm per pixel. Prior to each experiment, a reference background image was created using distilled water, which was then discarded. A volume of 3.5 ml of the suspension of Microcystis colonies was pipetted in between the shear surfaces. The shear protocol was as follows: pre-shear at ε˙=1.1m2/s3 for 1 minute to mix the suspension; 12 shear cycles of 5 minutes each followed by 12 cycles of 20 minutes each at the target dissipation rate (imaging cycles, at moderate dissipation rate of 0.019 m2/s3 for 1 min, were intercalated with each shear cycle). During the imaging cycles, 200 frames were captured at a rate of 10 fps, ensuring that each frame was sampled over a different set of colonies. Examples of the captured frames are shown in Figure 1.

D. Image analysis

Earlier version: The frames underwent pre-processing with background subtraction and thresholding (SI Appendix). Image segmentation was performed with the Scikit Image library for Python [80]. Colony size was determined by their Feret diameter d, which represents the maximum distance between two boundary points of the colony. The number of cells in a colony was estimated from the diameter using the formula Nc=(d/d1)f , where f is the average fractal dimension of the colonies. Since not all colonies were within the depth-of-focus, a correction function for the colony count was estimated using a suspension with a known size distribution. The image magnification was optimized for intermediate and large colonies ,given the wide range of colonies (from single cells ≈ 7μm to large colonies ≈ 500μm). Colonies smaller than 25 μm were measured with lower precision, necessitating a second correction function for accurate counting (SI Appendix, Eq. S1). Typical uncertainties for the biovolume distribution are displayed in Figure S8. The uncertainty is higher for small colonies due to the uncertainty in image resolution, and also higher for very large colonies due to the low counts.

Revised version: The frames underwent pre-processing with background subtraction and thresholding (SI Appendix). Image segmentation was performed with the Scikit Image library for Python [86]. Colony size was determined by their Feret diameter d, which represents the maximum distance between two boundary points of the colony. The number of cells in a colony was estimated from the diameter using the formula Nc = (d/d1)f, where f is the average fractal dimension of the colonies. Since not all colonies were within the depth-of-focus, a correction function for the colony count was estimated using a suspension with a known size distribution. The image magnification was optimized for intermediate and large colonies, given the wide range of colonies (from single cells ≈ 7μm to large colonies ≈ 500μm). Colonies smaller than 25μm were measured with lower precision, necessitating a second correction function for accurate counting (SI Appendix, Eq. S1). Typical uncertainties for the biovolume distribution are displayed in Figure S8. The uncertainty is higher for small colonies due to the uncertainty in image resolution, and also higher for very large colonies due to the low counts.

Earlier version: To estimate the colony fractal dimension f, a reference sample of colonies had its size distribution characterized in a cylindrical counting chamber, followed by the disaggregation into single cells and the measurement of cell concentration in a haemocytometer. The fractal dimension was computed by fitting the equation, cV=∑klkf , where c is the cell concentration, and lk is the relative diameter of each colony in the measured volume V. For Microcystis strain V163, a value of f = 2.09 was obtained, while for the field samples of Microcystis spp. it was 1.64. At least 3000 colonies were counted for each sample, giving a maximum uncertainty of δf < 0.04

Revised version: To estimate the colony fractal dimension f , a reference sample of colonies had its size distribution characterized in a cylindrical counting chamber, followed by the disaggregation into single cells and the measurement of cell concentration in a haemocytometer. The fractal dimension was computed by fitting the equation, cV=∑κlκf , where c is the cell concentration, and lk is the relative diameter of each colony in the measured volume V . For Microcystis strain V163, a value of f = 2.09 was obtained, while for the field samples of Microcystis spp. it was 1.64. At least 3000 colonies were counted for each sample, giving a maximum uncertainty of δf < 0.04

For the direct computation of the fragmentation frequency, we assume a pure fragmentation with an ideal binary erosion, such that the time derivative of the colony size distribution is ddtN[i](t)=−κ[i]N[i](t)lf[i]−lf[i−1]+κ[i+1]N[i+1](t)lf[i+1]−lf[i], (12)

where N [i](t), κ [i], and l [i] refer to the concentration distribution (number of counted colonies per volume per bin width), fragmentation frequency, and relative diameter at the center of bin i. This formula is only valid for the bins larger than the single cell, as the latter does not suffer fragmentation. The concentration distribution N [i] is initialized with the experimental value for t = 0, and an initial guess for the fragmentation frequency κ[i] is made for all bins larger than a single cell. The values of N[i](t) are evolved in time using an implicit Radau scheme. For each bin, the cumulative squared error between the experimental and numerical values of N[i](t) are computed and the error is minimized using the values of κ[i] as the optimization parameters. The uncertainty of κ[i] is propagated from the uncertainty of the concentration distribution. The script for image analysis is available here. Detected features for all datasets and raw images for the dataset in Fig. 2A-E are available here.

E. Yield stress of cyanobacterial colonies

Earlier version: The mechanical resistance of Microcystis colonies can be estimated by separating the colonies from the medium and measuring the yield stress of the concentrated aggregate. A suspension of Microcystis strain V163 was filtered through a 53μm mesh, and the large colonies were resuspended in fresh BG-11 medium. The sample was centrifuged for 30 min at 1000 g , and the pellet was separated. A dense gelatinous phase composed of cells and EPS was obtained, with cells accounting for 12 % of the total biovolume fraction. The yield stress of the concentrated aggregate was measured with an Anton Paar MCR 301 rheometer. The upper probe was a sand-blasted 50 mm cone with 1° angle, while the lower surface was covered with sandpaper to avoid wall slip. Each sample was pre-sheared at 10 s−1 for 2 minutes. Steady-shear and large amplitude oscillatory shear tests were performed (SI Appendix, Fig. S5). From a Herschel-Bulkley fit over the steady shear data, τ=τy+kγ˙n , the dynamical yield stress τy=4.3±0.3Pa (SD from best fit) was estimated. Due to the limited volume of samples available, only one repetition was possible for steady shear and oscillatory shear tests each.

Revised version: The mechanical resistance of Microcystis colonies can be estimated by separating the colonies from the medium and measuring the yield stress of the concentrated aggregate. A suspension of Microcystis strain V163 was filtered through a 53 μm mesh, and the large colonies were resuspended in fresh BG-11 medium. The sample was centrifuged for 30 min at 1000 g, and the pellet was separated. A dense gelatinous phase composed of cells and EPS was obtained, with cells accounting for 12% of the total biovolume fraction. The yield stress of the concentrated aggregate was measured with an Anton Paar MCR 301 rheometer. The upper probe was a sand-blasted 50 mm cone with 1 ° angle, while the lower surface was covered with sandpaper to avoid wall slip. Each sample was pre-sheared at 10 s-1 for 2 minutes. Steady-shear and large amplitude oscillatory shear tests were performed (SI Appendix, Fig. S5). From a Herschel-Bulkley fit over the steady shear data, τ=τy+κγ˙n , the dynamical yield stress τy = 4.3 ± 0.3 Pa (SD from best fit) was estimated. Due to the limited volume of samples available, only one repetition was possible for steady shear and oscillatory shear tests each.

F. Simulations

Earlier version: The balance equation (1) for the total normalized biovolume distribution was solved numerically using a zeroorder discrete population balance, or sectional method, in which the size distribution is described by uniformly sized bins of width Δl = 0.5, and the biovolume distribution is evaluated at a representative size centered on the bin[48]. The integrals for the aggregation (Eq. 2), fragmentation (Eq. 4), and transfer (Eq. 6) rates were discretized with the approach by [49]. The biovolume distribution evolved in time with a Runge-Kutta 4-5 adaptive time step scheme written in Fortran. The simulation results were fitted with the experimental data by minimizing the absolute error of each category’s 25 th, 50 th, and 75 th percentiles of the biovolume distribution, using a Nelder-Mead algorithm [81]. The stickiness α1 and for the pair of constitutive parameters S1 and q1 for the critical size of small colonies (category C1) were obtained by fitting the model to the dataset of a single-cell suspension under ε˙=0.019m2/s3 and ϕ = 0.01. The pair of constitutive parameters S2 and q2 for the critical size of large colonies (category C2) were obtained by fitting the model to the dataset of a large colony suspension under ε˙=5.8m2/s3 and ϕ = 0.01. These datasets were chosen as their aggregation/fragmentation times were similar to the measurement time, while for the other datasets, the variations were either too fast or slow to be captured in detail. In our simulations, the stickiness parameter α1 controls the rate at which an initial single cell suspension achieves an aggregated steady state. The parameters Si are the dissipation rates for which all colonies in category i are broken, and a higher Si shifts the steady state colony size upward. Finally, the exponents qi controls the sensitive of the fragmentation frequency to colony size. A high value of qi, as observed in our experiments, implies a low sensitivity to size and leads to a broader steady state size distribution. The script for the simulation is available here.

Revised version: The balance equation (1) for the total normalized biovolume distribution was solved numerically using a zeroorder discrete population balance, or sectional method, in which the size distribution is described by uniformly sized bins of width Δ l = 0.5, and the biovolume distribution is evaluated at a representative size centered on the bin[55]. The integrals for the aggregation (Eq. 2), fragmentation (Eq. 4), and transfer (Eq. 6) rates were discretized with the approach by [56]. The biovolume distribution evolved in time with a Runge-Kutta 4-5 adaptive time step scheme written in Fortran. The simulation results were fitted with the experimental data by minimizing the absolute error of each category’s 25th, 50th, and 75th percentiles of the biovolume distribution, using a Nelder-Mead algorithm [87]. The stickiness α1 and for the pair of constitutive parameters S1 and q1 for the critical size of small colonies (category C1) were obtained by fitting the model to the dataset of a single-cell suspension under ε˙=0.019m2/s3 and ϕ = 0.01. The pair of constitutive parameters S2 and q2 for the critical size of large colonies (category C2) were obtained by fitting the model to the dataset of a large colony suspension under ε˙=5.8m2/s3 and ϕ = 0.01. These datasets were chosen as their aggregation/fragmentation times were similar to the measurement time, while for the other datasets, the variations were either too fast or slow to be captured in detail. In our simulations, the stickiness parameter α1 controls the rate at which an initial single cell suspension achieves an aggregated steady state. The parameters Si are the dissipation rates for which all colonies in category i are broken, and a higher Si shifts the steady state colony size upward. Finally, the exponents qi controls the sensitive of the fragmentation frequency to colony size. A high value of qi , as observed in our experiments, implies a low sensitivity to size and leads to a broader steady state size distribution. The script for the simulation is available here.